```scala
import $file.hw2stdlib_new
import hw2stdlib_new._
```

Output:
```
Compiling hw2stdlib_new.sc
[32mimport [39m[36m$file.$            
[39m
[32mimport [39m[36mhw2stdlib_new._[39m
```

# Homework 2

Due 9/12 at 11:59pm

## Submission Instructions
Upload only this .ipynb file to Canvas. Do not add anything to hw2stdlib since you can't submit it.

In this homework we will develop more useful functions for numbers as well as some of the functions on lists that we spoke about during lecture.

## Arithmetic Functions

### Problem 1 (5 Points)

Write the minus function, which subtracts a natural number from another natural number. If the result would be less than zero then leave the answer as zero.


<!-- grade_id: cell-19061b6ebc83ee13, points: 0, grade: false, solution: true -->
```scala
def minus(x : Nat, y : Nat) : Nat =
    // BEGIN SOLUTION
    (x, y) match {
        case (Zero, _)          => Zero
        case (x, Zero)          => x
        case (Succ(x), Succ(y)) => minus(x, y)
    }
    // END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mminus[39m
```

<!-- grade_id: cell-80b015bfaabd7c53, points: 3, grade: true, solution: false -->
```scala
assert(minus(four, two) == two, 1)
assert(minus(four, Zero) == four, 2)
assert(minus(Zero, four) == Zero, 3)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-3fb81b3d6c3bdcd6, points: 2, grade: true, solution: false -->
```scala
// HIDDEN TEST (2 pts)
// BEGIN HIDDEN TESTS
assert(minus(Zero, Zero) == Zero, 1)
// END HIDDEN TESTS
```

### Problem 2 (5 points)

Define the $\texttt{pow}$ function, which raises a natural number to some power. Notationally: 

$$
f\ x\ y = x^y
$$

Hint: Use the mult function from the stdlib for help

<!-- grade_id: cell-cf03509753bbf640, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def pow(x : Nat, y : Nat) : Nat = y match {
    case Zero       => one
    case Succ(Zero) => x
    case Succ(yy)   => mult(x, pow(x, yy))
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mpow[39m
```

<!-- grade_id: cell-cc0a85f98f9550b1, points: 4, grade: true, solution: false -->
```scala
assert(pow(three, two) == nine, 1)
assert(pow(four, Zero) == one, 2)
assert(pow(two, two) == four, 3)
passed(4)
```

Output:
```
*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-e4f7652875d80790, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
// BEGIN HIDDEN TESTS
assert(pow(two, one) == two, 3)
// END HIDDEN TESTS
```

### Problem 3 (5 points)

Write the "equals to" function for numbers. This function should return $True$ if $x == y$ and $False$ otherwise.

<!-- grade_id: cell-5c44d445fb4cbec0, points: 0, grade: false, solution: true -->
```scala
def eq_nat(x : Nat, y : Nat) : Bool =
    // BEGIN SOLUTION
    (x, y) match {
        case (Zero, Zero)       => True
        case (Succ(x), Succ(y)) => eq_nat(x, y)
        case (x, y)             => False 
    }
    // END SOLUTION
```

Output:
```
defined [32mfunction[39m [36meq_nat[39m
```

<!-- grade_id: cell-cc0a85f98f9550b2, points: 4, grade: true, solution: false -->
```scala
assert(eq_nat(four, four) == True, 1)
assert(eq_nat(four, Zero) == False, 2)
assert(eq_nat(eight, nine) == False, 3)
assert(eq_nat(nine, nine) == True, 4)
passed(4)
```

Output:
```
*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-e4f7652875d80791, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
// BEGIN HIDDEN TESTS
assert(eq_nat(Zero, Zero) == True, 1)
// END HIDDEN TESTS
```

## Functions on Number Lists

### Problem 4 (5 pts)

Write the sum function. This function should take a list of natural numbers and return their sum. Here is the type:

$$
sum : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have sum 0.

<!-- grade_id: cell-1026d25f142fc212, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def sum(xs : List[Nat]) : Nat = xs match {
    case Empty       => Zero
    case Cons(x, xs) => plus(x, sum(xs))
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36msum[39m
```

<!-- grade_id: cell-c8396ec50b153d54, points: 4, grade: true, solution: false -->
```scala
assert(sum(Cons(one, Empty)) == one, 1)
assert(sum(Cons(one, Cons(seven, Cons(two, Empty)))) == ten, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-d118cae84cab07ee, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
// BEGIN HIDDEN TESTS
assert(sum(Empty) == Zero)
// END HIDDEN TESTS
```

### Problem 5 (5 pts)

Write the product function. This function should take alist of natural numbers and return their product. Here is the type:

$$
prod : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have product 1.

<!-- grade_id: cell-6a2caa631bc53781, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def prod(xs : List[Nat]) : Nat = xs match {
    case Empty       => one
    case Cons(x, xs) => mult(x, prod(xs))
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36mprod[39m
```

<!-- grade_id: cell-cc0a85f98f9550b3, points: 4, grade: true, solution: false -->
```scala
assert(prod(Cons(one, Empty)) == one, 1)
val fourteen = Succ(Succ(Succ(Succ(ten))))
assert(prod(Cons(one, Cons(seven, Cons(two, Empty)))) == fourteen, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
[36mfourteen[39m: [32mSucc[39m = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))))))))
```

<!-- grade_id: cell-e4f7652875d80792, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
// BEGIN HIDDEN TESTS
assert(prod(Empty) == one)
// END HIDDEN TESTS
```

### Problem 6 (6 pts)

Write the function that tells us if a given natural number is in a list of natural numbers. (you will need to use the equals function defined in problem 3)

$$
in : \mathbb{N} \rightarrow \text{List}\ \mathbb{N} \rightarrow \mathbb{B}
$$

<!-- grade_id: cell-603f77bab9b8f27a, points: 0, grade: false, solution: true -->
```scala
// BEGIN SOLUTION
def in(n : Nat, xs : List[Nat]) : Bool = xs match {
    case Empty       => False
    case Cons(x, xs) => eq_nat(x, n) match {
        case True  => True
        case False => in(n, xs)
    }
}
// END SOLUTION
```

Output:
```
defined [32mfunction[39m [36min[39m
```

<!-- grade_id: cell-b9d72e55abc06dd4, points: 6, grade: true, solution: false -->
```scala
assert(in(one, Cons(one, Empty)) == True, 1)
assert(in(seven, Cons(one, Cons(seven, Cons(two, Empty)))) == True, 2)
assert(in(eight, Cons(one, Cons(seven, Cons(two, Empty)))) == False, 3)
assert(in(seven, Empty) == False, 4)
passed(6)
```

Output:
```
*** Tests Passed (6 points) ***
```